You should consider the following before turning this problem in: Make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

### Notebook Assignment Module 3

# Optional Exercises

This Assignment is part of Module 3 of the course GIS222 (Fall2019).

For this problem set, we'll be using Jupyter notebook: [<img src="img/jupyter.png" alt="Jupyter Project Logo" title="Jupyter Project" width="100" />](https://jupyter-notebook.readthedocs.io/en/stable/)

Enter your solution where it says **YOUR CODE HERE** or **YOUR ANSWER HERE** and, if remaining, **delete** the statement **"raise NotImplementedError()"** when filling in your solution. For some tasks, additional interactive testing cells allow you to test your code.

Fill also your name and collaborators below.

```python
NAME = "William Ryan"
COLLABORATORS = ""
```

## Part A: Python Objects (3 points)

What does it mean if an object is mutable? What about immutable? Name at least two mutable and two immutable objects in Python. 

<!-- grade_id: MutableObjects, points: 3, grade: true, solution: true -->
Simply put, mutable means an object can be changed after it is created. Lists and sets are mutable. On the other hand immutable means unable to be changed. This includes Strings and Tuples.

---
## Part B: Temperature Conversion (2 points)

For this task, let's the code the algorithm of the temperature conversion, which you have had forumulated in pseudocode during the course module 1.

Complete a Python function that is supposed convert the input Fahrenheit degree into its Celsius degree equivalent, based on the following equation: 

$C=5/9 * (F-32)$

For that translate the equation into code, and enter it to replace the commented line (`### YOUR CODE HERE`). Make sure to use the variable names `C` and `F` for the two temperature variables and keep the indent! Also leave lines unchanged that contain such the comment `leave this line unchanged`.

Note: Commented parts in the code are ignored by the interpreter. It starts with a `#` and ends at the end of a line. The Jupyter Notebook colors comments in an olive green color.



<!-- grade_id: cell-Banswer, points: 0, grade: false, solution: true -->
```python
""" Enter the conversion equation below, keep the indent and 
    use C and F for the temperature values!!!
    Do not change the lines that state 'leave this line unchanged'  """

def tempConvert(F):   # leave this line unchanged
    C = 5/9*(F-32)
    
    return(C)         # leave this line unchanged
```

In the next code cell, assign a variable `tempF` with any value for a Temperature in Degree-Fahrenheit. Then run the two cells further below to receive converted value to Degree-Celcius and test your code. You can try different values. If you are satisfied with the test, move on to the next task.

<!-- grade_id: cell-Banswer2, points: 0, grade: false, solution: true -->
```python
"""Assign the variable 'tempF' with any value for a Temperature in Degree Fahrenheit."""

tempF = 32

```

```python
"""Check that your code returns the converted Temperature. """
 
tempConvert(tempF)
```

Output:
```
0.0
```

<!-- grade_id: cell-Bcheck2, points: 2, grade: true, solution: false -->
```python
"""Check that your assigns a number to tempF and that it returns a correctly converted Temperature. """

assert type(tempF) == float or int
assert tempConvert(32) == 0
assert tempConvert(50) == 10
assert tempConvert(100) == 37.77777777777778


```

---
## Part C:  Boolean Objects (2 Points)

Assign the variable a, b and c with number or string values, so that all tests in the following cell return True.

<!-- grade_id: cell-Canswer, points: 0, grade: false, solution: true -->
```python
""" Assign the variable a, b and c with number or string values,
    so that all tests in the following cell return True """

c=0
b=0
a='spam'
```

```python
"""Check that your code returns only True test results"""

Test1 = (type(c) == int)
Test2 = not b and bool(a)
Test3 = not(c) and a == 'spam'
Test4 = (type(c) == type(b))

print(Test1, Test2, Test3, Test4)
```

Output:
```
True True True True
```

<!-- grade_id: cell-Ctest, points: 2, grade: true, solution: false -->
```python
"""Check that your code returns only True test results"""
assert Test1 == True
assert Test2 == True
assert Test3 == True
assert Test4 == True
assert type(a) == int or float or str
assert type(b) == int or float or str
assert type(c) == int or float or str
```

---

## Part D: Sequence operations with Strings and Lists (3 points)
This tasks includes several string-list operations. You should process a string variable with the name `strIn` in the following way:

Split the string variable `strIn` into a list of strings with a space delimiter. Use the method `split()`
Then, using only list comprehensions, add % in front and at the end of every single word in the string list `strIn` (from previous task item). Assign the result to a new variable of the name `strListOut`.

Add your code sequence in the following code block where it says YOUR CODE HERE. It will complete the following code block (function). You may code this in one or two steps. But, the variables for input (`strIn`) and output (`strListOut`) of your code sequence have to have the same names as stated in the task.

Example: 

A string `"Spring is Awesome in Tempe"` would become a list like this: `['%Spring%','%is%', '%awesome%', '%in%', '%Tempe%']`


<!-- grade_id: cell-Danswer1, points: 0, grade: false, solution: true -->
```python
def stringManipulation(strIn): # leave this line unchanged
    
    """Enter your answer below, do not change the indent!"""
    splitList = strIn.split()
    strListOut = []
    for Word in splitList:
        strListOut.append('%' + Word + '%')


    return(strListOut) # leave this line unchanged
```

```python
"""Check that your code correctly manipulated the input string. 
   Use any other non-empty string as example."""

stringExample = "Is there anybody out there?"
stringManipulation(stringExample)
```

Output:
```
['%Is%', '%there%', '%anybody%', '%out%', '%there?%']
```

<!-- grade_id: cell-Xtest, points: 3, grade: true, solution: false -->
```python
"""Check that your code correctly manipulates the input string."""

stringExample1 = "Spring is Awesome in Tempe"
assert stringManipulation(stringExample1) == ['%Spring%', '%is%', '%Awesome%', '%in%', '%Tempe%']
assert stringManipulation(stringExample1)[0][0] == '%'
assert stringManipulation(stringExample1)[0][-1] == '%'
assert type(stringManipulation(stringExample1)) == list

```

---
## Part E: List Comprehension with Number Lists

### Task E.1 (3 points)

This tasks includes writing list comprehensions and other operations:

* Write a a single list comprehension to print all the numbers in the range 1-100 that are either divisible by 10 or 11. Assign the result to a list `outList`. Use the built-in function `range()`.


Enter the code for these task items in the cell below. Separate the blocks of the tasks items with a commented line,  starting with `# `. Use exactly these variable names, put your code below, where it states `YOUR CODE HERE` and keep the indent as provided.

<!-- grade_id: cell-E1answer, points: 0, grade: false, solution: true -->
```python
"""Write a list comprehension to find numbers divisible 
   by at least one of two given integer numbers. """

def divisionList(a,b): # leave this line unchanged
    
    """Enter your answer below, do not change the indent!"""
    outList = []
    for x in range (1, 101):
        if x%a == 0 or x%b == 0:
            outList.append(x)
    

    return(outList) # leave this line unchanged

```

```python
"""Adjust the two integer numbers to your liking."""
a = 15
b = 21

```

```python
"""Check the result of your code."""
print("The following numbers between 1 and 100 are divisable by {} or {}: ".format(a,b))
print(divisionList(a,b))
print("\n")
```

Output:
```
The following numbers between 1 and 100 are divisable by 15 or 21: 
[15, 21, 30, 42, 45, 60, 63, 75, 84, 90]
```

<!-- grade_id: cell-E1test, points: 3, grade: true, solution: false -->
```python
"""Check the result of your code."""

assert divisionList(99,100) == [99, 100]
assert divisionList(15,21)  == [15, 21, 30, 42, 45, 60, 63, 75, 84, 90]


```

### Task E.2

You have given two lists containing a student grade book. Below that, samples show you more sophisticated list comprehensions. Study those examples and experiment with them in the available cells, further below.

```python
# Sample grade book data
#             0      1       2        3      4      5       6        7         8
headings = ["ID", "Name", "Gender", "HW1", "HW2", "HW3", "Exam1", "Exam2", "FinalExam"]

grades = [
#    0          1          2    3   4   5   6   7   8
 ["b0001", "Jane Doe",    "F", 95, 87, 92, 88, 93, 90] # 0
,["b0002", "John Blue",   "M", 55, 76, 89, 77, 82, 80] # 1
,["b0003", "Kim Tester",  "F", 80, 75, 65, 70, 75, 80] # 2
,["b0004", "Larry Black", "M", 90, 90, 92, 95, 85, 94] # 3
,["b0005", "Susan White", "F", 65, 52, 85, 45, 80, 82] # 4
]

```

```python
# samples that use the Sample Grade Book in sample_grade_book.py
# Don't import this -- just use cut-and-paste to test the individual
#    expressions.

# display the gradebook (will need to type this one in 2 lines)
for g in grades:
    print(g)

# Larry Black's grade on Exam2
grades[3][7]

# Suppose you don't know the row/column numbers.
# finding Larry Black's row
headings.index('Name') # need to know the name column first

[g[headings.index('Name')] for g in grades] # names of all students

[g[headings.index('Name')] for g in grades].index('Larry Black') # find Larry Black's index

# to get Larry Black's complete row in the gradebook
grades[[g[headings.index('Name')] for g in grades].index('Larry Black')]

# to get Larry Black's grades
grades[[g[headings.index('Name')] for g in grades].index('Larry Black')][headings.index('HW1'):]

#finding Exam2's column
headings.index('Exam2')

#finding Larry Black's grade on Exam2
grades[[g[headings.index('Name')] for g in grades].index('Larry Black')][headings.index('Exam2')]

# grades on Exam2
[g[7] for g in grades]
[g[headings.index('Exam2')] for g in grades]

# average on hw2
hw2 = [g[4] for g in grades]
float(sum(hw2))/len(grades)
# average on Exam2
float(sum([g[headings.index('Exam2')] for g in grades]))/len(grades)


# highest average
h = max([sum(g[3:])/float(len(g[3:])) for g in grades])
# index of that person
pindex = [sum(g[3:])/float(len(g[3:])) for g in grades].index(h)
# name and average
grades[pindex][1], h
# all together
grades[[sum(g[3:])/float(len(g[3:])) for g in grades].index(max([sum(g[3:])/float(len(g[3:])) for g in grades]))
][1], max([sum(g[3:])/float(len(g[3:])) for g in grades])
```

Output:
```
['b0001', 'Jane Doe', 'F', 95, 87, 92, 88, 93, 90]
['b0002', 'John Blue', 'M', 55, 76, 89, 77, 82, 80]
['b0003', 'Kim Tester', 'F', 80, 75, 65, 70, 75, 80]
['b0004', 'Larry Black', 'M', 90, 90, 92, 95, 85, 94]
['b0005', 'Susan White', 'F', 65, 52, 85, 45, 80, 82]
('Larry Black', 91.0)
```

Try the examples by copying & pasting them into the code cells below. Study their output. You may also add more code cells, if you like. This exercise will not be graded, but it is helpful to understand list comprehensions at a very deep level.

```python
### Try the examples in this code cell. (This will not be graded.)


```

```python
### Try the examples in this code cell. (This will not be graded.)


```

```python
### Try the examples in this code cell. (This will not be graded.)


```

```python
### Try the examples in this code cell. (This will not be graded.)


```

---
## Part F: Dictionaries 

### Task F.1  (3 points)
For this task, generate a new dictionary that merges the lists `headings` and `grades` from *Task E.2* into one dictionary of the name `students`. 
The dictionary `students` should be a nested dictionary: a dictionary of 5 dictionaries, with each of these 5 dictionaries containing a row of the list `grades`. Use the student ID's as key of the outer dictionary. Then use the remaining header lines as keys of the inner dictionaries and the remaing columns as values of the inner dictionaries. It is sufficient to include only one of the grade columns, e.g. `HW1` or `Exam2`.
For the dictionary keys, use the name names as listed in `headings` above. Use appropriate object types for the content numbers and strings.

<!-- grade_id: cell-F1answer, points: 2, grade: true, solution: true -->
```python
"""Convert the list grades into a dictionary of the name students."""

students = {'b0001': {'Name': 'Jane Doe', 'Gender': 'F', 'HW1': 95, 'Exam1': 88, 'FinalExam': 90},
    'b0002': {'Name': 'John Blue', 'Gender': 'M', 'HW1': 55, 'Exam1': 77, 'FinalExam': 80},
    'b0003': {'Name': 'Kim Tester', 'Gender': 'F', 'HW1': 80, 'Exam1': 70, 'FinalExam': 80},
    'b0004': {'Name': 'Larry Black', 'Gender': 'M', 'HW1': 90, 'Exam1': 95, 'FinalExam': 94},
    'b0005': {'Name': 'Susan White', 'Gender': 'F', 'HW1': 65, 'Exam1': 45, 'FinalExam': 82}}
```

<!-- grade_id: cell-F1test, points: 1, grade: true, solution: false -->
```python
"""Check the result of your code."""

assert type(students) == dict
assert list(students.keys())[0] == 'b0001'
assert len(students) == 5
assert len(students['b0001']) > 2
assert students['b0001']['Name'] == 'Jane Doe'
assert students['b0002']['Gender'] == 'M'

```

### Task F.2  (1 point)
Now, given this dictionary, how can you look up a any student's (e.g. Larry Black's) grade in any exam (e.g. Exam2) from the dictionary students? Write a chained list comprehension in the code cell below. 

<!-- grade_id: cell-F2answer, points: 1, grade: true, solution: true -->
```python
"""Enter your code."""

students['b0004']

# Check the result of your code, below
```

Output:
```
{'Name': 'Larry Black', 'Gender': 'M', 'HW1': 90, 'Exam1': 95, 'FinalExam': 94}
```

---
## Part G: Ungraded Tasks for Further Practise.

In this part, a few more coding exercises are given, which you may use for practise. However, these will not be tested or graded.

* Write a list comprehension, that generates a list of numbers between 1 and 1000 in steps of 100 and then returns the square of each value. (Hint: `range()`).
* Using list comprehensions only, create a 4x4 idendity matrix.
* Develop the code that transfers a list, L, of integers into a second list L2 containing the only the even numbers from L. (Hint: Modulus Operator)
* Using the student gradebook dictionary form above, write a list comprehension (chain), which can find the Person with the highest grade in Exam 2, without knowing the content (but the structure) of the dataset.

```python

```

```python

```

```python

```

```python

```

```python

```

If you are satisfied with the test, validate your notebook, save it and make a backup until you submit this assignment.